# CS431 - CÁC KỸ THUẬT HỌC SÂU VÀ ỨNG DỤNG

**Hướng dẫn**:
- Hoàn tất tất cả các mục có comment `YOUR CODE HERE`


---

# LUYỆN TẬP NUMPY - PHẦN 2

```python
import numpy as np
```

## Bài 1: Gửi tiết kiệm
Tính số tiền còn lại trong tài khoản ngân hàng với đầu vào là
danh sách các số tiền giao dịch, tương ứng là sẽ có danh sách
các loại giao dịch gửi tiền (ký hiệu bằng chuỗi 'D'), rút tiền
(ký hiệu bằng chuỗi 'W')

**Input:**
   - moneys: kiểu numpy array, chứa danh sách các số tiền giao dịch
   (gửi và rút khỏi tài khoản)
   - types: kiểu numpy array, chứa danh sách các loại giao dịch
   (gửi ký hiệu là 'D', rút ký hiệu là 'W')

**Output:**
   - Kiểu số thực, số tiền còn lại trong tài khoản

Ví dụ:

Input:
```
moneys = [400.0, 100.0, 200.0]
types = ['D', 'W', 'D']
```

Output:
```
500.0
```

<!-- grade_id: cell-048577ccfe8e62c6, points: 0, grade: false, solution: true -->
```python
def balance(moneys, types):
    remaining_balance = 0.0
    
    # Duyệt qua từng giao dịch và loại giao dịch tương ứng
    for money, transaction_type in zip(moneys, types):
        if transaction_type == 'D':
            # Nếu là giao dịch gửi tiền ('D'), thì cộng vào số tiền còn lại
            remaining_balance += money
        elif transaction_type == 'W':
            # Nếu là giao dịch rút tiền ('W'), thì trừ đi từ số tiền còn lại
            remaining_balance -= money
    
    return remaining_balance
```

```python
# thử nghiệm với ví dụ của đề bài
balance(np.array([400.0, 100.0, 200.0]), np.array(['D', 'W', 'D']))
```

Output:
```
500.0
```

<!-- grade_id: cell-e03038c165ddaafc, points: 5, grade: true, solution: false -->
```python
# cell chấm điểm 1
assert balance(np.array([400.0, 100.0, 200.0]), np.array(['D', 'W', 'D'])) == 500
```

<!-- grade_id: cell-839d4d8b7de04b4c, points: 5, grade: true, solution: false -->
```python
# cell chấm điểm 2
```

## Bài 2: Điểm trung bình
Tính điểm trung bình của một bạn sinh viên biết trước danh sách các điểm số và số tín chỉ của từng môn học.

**Input:**
   - marks: kiểu List, gồm *n* phần tử. Đây là điểm các môn của một sinh viên.
   - credits: kiểu List, gồm *n* phần tử. Đây là số tín chỉ lần lượt của các môn.   
   
**Output:**
   - Kiểu số thực, điểm trung bình của các môn học của sinh viên

Ví dụ:

Input:
```
marks = [8.0, 9.0, 10.0]
credits = [2, 2, 1]
```

Output:
```
8.8
```

<!-- grade_id: cell-a0d93b1f727b2f51, points: 0, grade: false, solution: true -->
```python
def gpa_one_student(marks, weights):
    result = float(np.dot(marks,weights))/float(np.sum(weights))
    return float(result)
```

```python
# thử nghiệm với ví dụ của đề bài
marks = [8.0, 9.0, 10.0]
credits = [2, 2, 1]
gpa_one_student(marks, credits)
```

Output:
```
8.8
```

<!-- grade_id: cell-f9ecf7150939ef80, points: 5, grade: true, solution: false -->
```python
# cell chấm điểm 3
assert abs(gpa_one_student([10, 9, 8], [4, 3, 2])-9.222222222222221) < 0.000001
```

<!-- grade_id: cell-4706026df476b116, points: 5, grade: true, solution: false -->
```python
# cell chấm điểm 4
```

## Bài 3:
Đếm số lượng sinh viên có điểm số không dưới K điểm cho trước của một lớp học.

**Input:**

    - marks: kiểu numpy array, kích thước 1 x n. Đây là điểm toán cả n sinh viên trong lớp.
    - K: kiểu số thực, điểm chuẩn.
   
**Output:**

    - Kiểu số nguyên, số lượng sinh viên trong lớp đạt từ K điểm trở lên

<!-- grade_id: cell-75bbb0d6e24b7c40, points: 0, grade: false, solution: true -->
```python
def pass_exam(marks, K):
    passed = marks > K
    return np.sum(passed)
    raise NotImplementedError()
```

```python
# Test thử với 1 mẫu ví dụ
marks = np.array([4.0, 9.0, 5.0])
K = 5
pass_exam(marks, K)
```

Output:
```
1
```

<!-- grade_id: cell-19c61f2fdfaa710a, points: 5, grade: true, solution: false -->
```python
# cell chấm điểm 5
assert pass_exam(np.array([3.0, 9.5, 4.5]), 5.0) == 1
```

<!-- grade_id: cell-6457e32df8fac035, points: 5, grade: true, solution: false -->
```python
# cell chấm điểm 6
```

## Bài 4:
Tính điểm trung bình cao nhất của một lớp học biết trước danh sách các điểm số và số tín chỉ
cho từng môn học của một lớp.

**Input:**
   - marks: kiểu numpy array, kích thước m x n. Đây là điểm các môn của một lớp học. Trong đó, m là số sinh viên, n là số môn học
   - credits: kiểu numpy array, kích thước 1 x n. Đây là số tín chỉ lần lượt của các môn.
   
**Output:**
   - Kiểu số thực, điểm trung bình của sinh viên cao nhất.

Ví dụ:

Input:
```
marks = np.array(
   [[8.0, 9.0, 10.0],
   [7.0, 9.0, 8.0],
   [8.0, 8.0, 8.0],
   [10.0, 9.0, 7.0],
   [9.0, 9.0, 10.0]
])

credits = np.array([2,2,1])
```

Output:
```
9.2
```

<!-- grade_id: cell-3a099ebd7eee51b7, points: 0, grade: false, solution: true -->
```python
def max_gpa(marks, credits):
    matrix = np.dot(marks,credits)
    return max(matrix/np.sum(credits))
    raise NotImplementedError()
```

```python
# thử nghiệm với ví dụ của đề bài
marks = np.array(
   [[8.0, 9.0, 10.0],
   [7.0, 9.0, 8.0],
   [8.0, 8.0, 8.0],
   [10.0, 9.0, 7.0],
   [9.0, 9.0, 10.0]
])
credits = np.array([2,2,1])
max_gpa(marks, credits)
```

Output:
```
9.2
```

<!-- grade_id: cell-84b3a393314e50ba, points: 5, grade: true, solution: false -->
```python
# cell chấm điểm 7
marks = np.array(
   [[4.0, 9.0, 4.5],
   [5.0, 8.0, 4.5],
   [6.0, 7.0, 7.5],
   [7.0, 6.0, 7.5],
   [8.0, 5.0, 4.5]
])
credits = np.array([2,1,2])
assert abs(max_gpa(marks, credits)-7.0) < 0.000001
```

<!-- grade_id: cell-afcc38c40ff1554e, points: 5, grade: true, solution: false -->
```python
# cell chấm điểm 8
```

## Bài 5:
Tính điểm trung bình tích luỹ của từng sinh viên trong lớp học biết trước
danh sách các điểm số và số tín chỉ cho từng môn học của một lớp.

**Lưu ý:** khi tính điểm trung bình ta sẽ loại bỏ các điểm và các tín chỉ của môn dưới 5 điểm (rớt).

**Công thức tính điểm trung bình tích lũy:**
   - ĐTBTL = (điểm môn 1 * số tín chỉ môn 1 + ... + điểm môn n * số tín chỉ môn n) / tổng số tín chỉ

**Input:**
   - marks: kiểu numpy array, kích thước m x n. Đây là điểm các môn của một lớp học.
   Trong đó, m là số sinh viên, n là số môn học
   - credits: kiểu numpy array, kích thước 1 x n. Đây là số tín chỉ lần lượt của các môn.

**Output:**
   - Kiểu numpy array, kích thước 1 x m, chứa điểm trung bình của các bạn sinh viên trong lớp.

Ví dụ:

Input:
```
marks = np.array([
    [8.0, 9.0, 10.0],
    [4.0, 9.0, 8.0],
    [8.0, 3.0, 8.0],
    [10.0, 9.0, 5.0],
    [9.0, 9.0, 4.0]
])

credits = np.array([2, 2, 1])
```

Output:
```
array([8.8, 8.66666667, 8., 8.6, 9.])
```

<!-- grade_id: cell-85c2823fb5e72e1e, points: 0, grade: false, solution: true -->
```python
def gpa_of_pass(marks, credits):
    valid_marks = np.where(marks >= 5, marks, 0)
    valid_credits = np.where(valid_marks > 0, credits, 0)

    total_points = np.sum(valid_marks * valid_credits, axis=1)
    total_credits = np.sum(valid_credits, axis=1)

    total_credits = np.where(total_credits == 0, 1, total_credits)

    return total_points / total_credits
```

```python
# thử nghiệm với ví dụ của đề bài
marks = np.array(
    [[8.0, 9.0, 10.0],
    [4.0, 9.0, 8.0],
    [8.0, 3.0, 8.0],
    [10.0, 9.0, 5.0],
    [9.0, 9.0, 4.0]])
credits = np.array([2,2,1])
gpa_of_pass(marks, credits)
```

Output:
```
array([8.8       , 8.66666667, 8.        , 8.6       , 9.        ])
```

<!-- grade_id: cell-fa2fa11d0d185e58, points: 5, grade: true, solution: false -->
```python
# cell chấm điểm 9
marks = np.array(
   [[4.0, 9.0, 4.5],
   [5.0, 8.0, 4.5],
   [6.0, 7.0, 7.5],
   [7.0, 6.0, 7.5],
   [8.0, 5.0, 4.5]])
credits = np.array([2,1,2])
assert np.sum(abs(gpa_of_pass(marks, credits)-np.array([9., 6., 6.8, 7.0, 7.0]))) < 0.000001
```

<!-- grade_id: cell-8119a54a9940d0cf, points: 5, grade: true, solution: false -->
```python
# cell chấm điểm 10
```

```python

```